<div class = "alert alert-warning" align="center"> 
    <img src="https://www.egd.mg/wp-content/uploads/2020/01/cropped-logo-home-110px.png" align="left" width="290"  margin="50px"> 
    <center><h1>Lycée français de Tananarive</h1> 
        <h2> Jupyter Hub,  Spécialité NSI </h2></center>
</div>

<div align=justify>
    <h1 style='text-align:center' ><span style = 'color:#0000FF'>NoteBook 1 : Quelques prérequis
        </span></h1>
    <h2 style='text-align:center' ><i>Résolution dichotomique</i></h2>
</div>

## 1. Passer une fonction en paramètre à une autre fonction.
Il existe plusieurs façons de passer en paramètre une fonction à une autre fonction en Python.

Première méthode on peut définir la fonction, par exemple définissons $f(x) = 2\cdot x^2 -x +1$

```python
def fonction(x) :
    return 2*x**2-x+1
```

Et ensuite il suffira de l'appeler dans une autre fonction par un `appliquer(fonction,3)` par exemple :

```python
def appliquer(f,x) :
    return f(x)
appliquer(fonction,3)
```

<div class = "alert alert-info">
    <b> Q1</b> : Créer une fonction nommée <code>fonction2(x)</code> qui permette d'implémenter la fonction :$f_2(x) = 4\cdot x^3+2\cdot x^2-3$
 </div>

<!-- grade_id: cell-39cc2e19abc4d876, points: 0, grade: false, solution: true -->
```python
def fonction2(x) :
    # YOUR CODE HERE
    raise NotImplementedError()
    print("")
```

<!-- grade_id: cell-a5cf477a9dafba43, points: 1, grade: true, solution: false -->
```python
assert appliquer(fonction2,1) == 3
assert appliquer(fonction2,9) == 3075
```

Mais on peut utiliser une autre solution avec la fonction $lambda$.

```python
appliquer(lambda x:2*x**2-x+1,4)
```

## 2. Le module math
Le module math met à votre disposition les fonctions $cos(x)$ et $\sqrt{x}$ entre autre :

```python
from math import *
cos(3.1)
```

```python
sqrt(29)
```

<div class = "alert alert-info">
    <b> Q2</b> : On a $f_3(x) = \sqrt{cos(2\cdot x)}$
    
Créer une variable <code>solution</code> et lui affecter la valeur $f_3(2.5)$ en utilisant la fonction <code>appliquer</code> en utilisant une fonction <code>lambda</code>
</div> 

<!-- grade_id: cell-e50aae767bf469ad, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b5c4de47bf8b990b, points: 1, grade: true, solution: false -->
```python
assert 0.53<solution<0.54
```

<div class="alert alert-warning" align="center">
    <b>Equipe pédagoqique LFT-NSI</b>
<br/>ressource éducative libre distribuée sous <a href="http://creativecommons.org/licenses/by-nc-sa/4.0/"!>Licence Creative Commons Attribution</a>
    <img src="https://licensebuttons.net/l/by-nc-sa/4.0/88x31.png" align="center" width="70"  margin="50px"> 
</div>