<div class = "alert alert-warning" align="center"> 
    <img src="https://www.egd.mg/wp-content/uploads/2020/01/cropped-logo-home-110px.png" align="left" width="290"  margin="50px"> 
    <center><h1>Lycée français de Tananarive</h1> 
        <h2> Jupyter Hub,  Spécialité NSI </h2></center>
</div>

<div align=justify>
    <h1 style='text-align:center' ><span style = 'color:#0000FF'>NoteBook 2 : Resoudre numériquement une équation
        </span></h1>
    <h2 style='text-align:center' ><i>Résolution dichotomique</i></h2>
</div>

Les méthodes numériques permettent de calculer de manière effectives des solutions numériques à divers problèmes, souvent liés à la physique.

Nous aborderons ici la méthode dichotomique pour résoudre une équation.

Lors de cette épreuve nous aurons besoin de la bibliothéque de fonctions mathématiques, commençons par l'importer

```python
from math import *
```

## 1. Le Sujet
La méthode dichotomique est utilisé dans la résolution de nombreux problèmes nuémriques. Elle permet, en particulier, de trouver la solution d'une équation $f(x)=0$ où $f$ est une fonction continue d'une variable réelle.

Cette méthode peut être utilisé dans les cas où une solution analytique ne peut être connue.

Dans toute la suite, on supposera que $f$ est une fonction continue strictement croissant ou décroissant sur l'intervalle considéré;

On admet que si il existe $a<b$ tels que $f(a)\cdot f(b)<0$ alors il existe $c \in [a,b]$ tel que $f(c)=0$.

<div class = "alert alert-info">
    <b>Q1</b> L'algorithme de recherche de solution par la méthode dichotomique consiste, à partir d'un intervalle $[a,b]$ contenant une solution $x$ à l'équation $f(x)=0$, à trouver un nouvel intervalle, deux fois plus petit, qui contient une solution.

Pour cela, on évalue $f(m)$ avec $m=\frac{a+b}{2}$. Alors au moins, un des deux intervalles $[a,m]$ ou $[m,b]$ contient une solution.
> Ecrire une fonction nommée `iteration_dicho` qui prend en praramètre la fonction $f$ , le deux valeurs  $a$ et $b$, et renvoie les bornes de l'intervalle, plus petit, contenant la solution.
</div>

<!-- grade_id: cell-a65d961fe4245a60, points: 0, grade: false, solution: true -->
```python
def iteration_dicho(f,a,b) :
    """ itération de fonction
    Prec :  f : fonction de R dans R continue sur l'intervalle [a,b]
            a et b : float, bornes min et max de l'intervalle contenant la solution
    Post : Un tuple contenant les bornes du nouvel intervalle. """
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
def g(x) :
    return x+1
def h(x) :
    return 5*x**5-4*x**3-x**2+3
```

<!-- grade_id: cell-53b1a266b5296d11, points: 0, grade: true, solution: false -->
```python
assert iteration_dicho(g,-2,3) == (-2, 0.5)
assert iteration_dicho(g,-3,0) == (-1.5, 0)
assert iteration_dicho(g,-1.5,0) == (-1.5, -0.75)
assert iteration_dicho(h,-3,0) == (-1.5, 0)
assert iteration_dicho(h,-1.5,-0.5) == (-1.5, -1.0)
assert iteration_dicho(h,-1.2,-1.1) == (-1.15, -1.1)
```

<div class = "alert alert-info">
<b>Q2</b> Il est très rare que la solution exacte soit trouvée, et tout en sachant qu'il faut tenir compte des erreurs de calcul sur les nombres flottants.

Nous allons donc considérer que si l'intervalle est suffisament petit ( $a-b$ inférieur à une valeur $\varepsilon$ fixée), alors une solution approchée est $\frac{a+b}{2}$, ce qui garanti une erreur inférieure à $\frac{\varepsilon}{2}$.

> Ecrire une fonction `dichotomie` qui prend en argument une fonction $f$, les bornes d'un intervalle $a$ et $b$ et la valeur de $\epsilon$ et renvoie la solution approchée de l'équation $f(x)=0$, ou `None`, si il n'existe pas à priori de solution dans l'intervalle considéré ( car $f(a) \cdot f(b) > 0$ par exemple, ce qui ne signifie pas d'assurer qu'il n'y a pas de racines dans l'intervalle considéré, mais ne permet pas non plus d'assurer qu'il y en a une).
</div>

<!-- grade_id: cell-43cf4e98893def32, points: 0, grade: false, solution: true -->
```python
def dichotomie(f,a,b,eps) :
    """ fonction dicho
    Prec :  f : fonction de R dans R continue sur l'intervalle [a,b]
            a et b : float, bornes min et max de l'intervalle contenant la solution
            eps : largeur de l'intervalle à atteindre
    Post : Un flottant représentant la solution approchée à eps/2 près """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-66747b0c7a8ed17a, points: 3, grade: true, solution: false -->
```python
assert dichotomie(g,-10,10,1e-3) == -1.00006103515625
assert dichotomie(g,-10,10,1e-8) == -1.0000000009313226
assert dichotomie(g,-10,10,1e-14) == -1.0000000000000009
assert dichotomie(h,-10,10,1e-2) == -1.0595703125
assert dichotomie(h,-10,10,1e-10) == -1.0574766924037249
```

<div class = "alert alert-info">
    <b>Q3</b> Tester votre fonction de résolution dichotomique sur un un cas connu..

Soit la fonction définie par : $f(x) = (x-1) \cdot (x+3)^3 \cdot (x-4)$. Cette fonction admet trois racines connues.
> Ecrire une fonction `resolution` qui renvoie les trois solutions approchées de cette fonction (et uniquement de cette fonction) sous forme d'un tuple avec une précision de $10^{-4}$.
 Pour vérifier la valeur approximative des racines, on peut tracer la courbe.
</div>

```python
import matplotlib.pyplot as plt
x = [i*0.01 for i in range(-500,500)]
y = [(i-1)*(i+3)**3*(i-4) for i in x]
y1 = [0 for i in x]

plt.plot(x,y)
plt.plot(x,y1)
```

<!-- grade_id: cell-c36a963576f43258, points: 0, grade: false, solution: true -->
```python
def resolution (f) :
    """ resolution de fonction
    Prec :  f : fonction (x-1)*(x+3)**3*(x-4)
    Post : Un tuple contenant les trois solutions approchées dans l'ordre croissant """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-4c707f3e0d5eb807, points: 2, grade: true, solution: false -->
```python
assert -3.00009<resolution(lambda x: (x-1)*(x+3)**3*(x-4))[0]<-2.99999
assert 0.9999<resolution(lambda x: (x-1)*(x+3)**3*(x-4))[1]<1.0001
assert 3.9999<resolution(lambda x: (x-1)*(x+3)**3*(x-4))[2]<4.0001
```

<div class = "alert alert-info">
    <b>Q4</b> Utiliser votre fonction <code>dichotomie</code> pour affecter à la variable <code>sol</code> la solution de l'équation $f(x)=cos(\sqrt{x})$ entre 10 et 30 avec une précision de $10^{-5}$ !
    
$\heartsuit $ : Pensez à ce que l'on a vu dans le Notebook précédent.
</div>

<!-- grade_id: cell-dc551d61c2d212e3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-521f04b3a7f1aeb4, points: 2, grade: true, solution: false -->
```python
assert 22<sol<23
```

<div class = "alert alert-info">
    <b>Q5</b> Le nombre d'étapes réalisées par l'algorithme dépend uniquement de la largeur de l'intervalle et de la borne $\varepsilon$ chosie.

En supposant $a=0$ et $b=8$ et $\epsilon = 10^{-2}$, combien de tour de boucles effectuera l'algorithme dans ce cas ?
</div>

<!-- grade_id: cell-06562ab66ff633aa, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

<div class = "alert alert-info">
    <b>Q6</b> En supposant que l'appel de fonction la question 5 s'effectue en 0.1 millisenconde, à combien peut on estimer le temps de calcul si la valeur de $\epsilon$ était de $10^{-6}$ ?
</div>

<!-- grade_id: cell-88e286c92ca1c1b9, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

<div class="alert alert-warning" align="center">
    <b>Equipe pédagoqique LFT-NSI</b>
<br/>ressource éducative libre distribuée sous <a href="http://creativecommons.org/licenses/by-nc-sa/4.0/"!>Licence Creative Commons Attribution</a>
    <img src="https://licensebuttons.net/l/by-nc-sa/4.0/88x31.png" align="center" width="70"  margin="50px"> 
</div>